# Visualisation with Python and Bokeh - Due 12th August (23:59)


You are required to produce a **visualisation** of food hygiene ratings for different establishments across London.
 
Your visualisation should contain a map-based display of the ratings, placed according to their geolocation data. A user should be able to intuitively see which businesses are ‘safe’ to eat at, and those which have not scored so well.
 
You will supplement the map with additional charts and graphics that you deem appropriate to tell a coherent story from the data available.

## Question 0

In this assignment, you will use Bokeh, in particular `WMTSTileSource` class to add points to a map.  You will use a convenience package [`bokeh.tile_providers`](http://bokeh.pydata.org/en/latest/docs/reference/tile_providers.html) which creates a `WMTSTileSource` (like used in the guided exercise) instance with the `url` and `attribution` already set.  So, instead of manually creating a tile, you can just use one of the variables already created.  See [the Bokeh source code](https://github.com/bokeh/bokeh/blob/master/bokeh/tile_providers.py) to look at how they have done it.

The available tiles supported by Bokeh use [_Web Mercator_](https://en.wikipedia.org/wiki/Web_Mercator) format to represent location, so a function `wgs84_to_web_mercator` to convert the two is provided.  

Run the code in the cell below to set up the Notebook.

```python
# You don't need to write anything here
# Set up MongoDB
from pymongo import MongoClient

client = MongoClient('mongodb://cpduser:M13pV5woDW@mongodb/health_data', 27017)
db = client.health_data

from nose.tools import *

# # Numpy, Pandas and Bokeh imports
import numpy as np
import pandas as pd
from bokeh.palettes import Spectral6
from bokeh.io import output_notebook, show
from bokeh.models.sources import ColumnDataSource
from bokeh.models import *
from bokeh.io import curdoc
from bokeh.tile_providers import *
from bokeh.models.tiles import WMTSTileSource
import ipywidgets
from ipywidgets import interact, interactive
from ipywidgets import HBox, Label, IntSlider

from bokeh.plotting import figure
from bokeh.models import TapTool, CustomJS




def wgs84_to_web_mercator(df, lon="lon", lat="lat"):
    """
    Converts decimal longitude/latitude to Web Mercator format
    Source https://github.com/bokeh/bokeh-notebooks/blob/master/tutorial/11%20-%20geo.ipynb
    """
    k = 6378137
    df["x"] = df[lon] * (k * np.pi/180.0)
    df["y"] = np.log(np.tan((90 + df[lat]) * np.pi/360.0)) * k
    return df

# from ipywidgets import *
# from bokeh.layouts import *
from IPython.display import display
from bokeh.io import output_file, output_notebook, show, push_notebook

output_notebook()
```

<!-- grade_id: setup, points: 0.0, grade: true, solution: false -->
```python
# You don't need to write anything here
# Check it's set up correctly
try:
    imports = [MongoClient, db, np, pd, output_notebook, show, ColumnDataSource,
               output_notebook, show, ColumnDataSource, STAMEN_TERRAIN, figure
              ]
    assert True
    print('Successfully imported required libraries')
except NameError as e:
    print(e)
    assert False
```

Output:
```
Successfully imported required libraries
```


## Question 1: Create Map

In this question, you will create functions which will **return** the different objects required for the visualisation of a map on: A [`DataFrame`](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.html), a `ColumnDataSource` and a `Figure`.

### Question 1(a) [4 marks]

Create a function **`get_data()`** to extract data from the MongoDB database for all institutions which are in the **London** region with the following constraints:
- The results should include: `Lat`, `Lng`, `BusinessType`, `AddressLine1`, `BusinessName`, `RatingValue` but **NOT** the `_id` field
- The results should **only include businesses which have a RatingValue** (N.B. A value of 0 is a RatingValue)
- The results returned should **only include businesses which have a Geocode**
- The returned values should be **limited to 200** institutions
- **Add fields `x` and `y` in _Web Mercator_ format** to specify co-ordinates on the map
- **Return** the result as as [`DataFrame`](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.html).  

*Hint: Week 3, Guided Exercise 4, Cursors*  
*Hint: Week 4, Guided Exercise 2, Importing Data*

```python
test = db.uk.find_one()
test.keys()
```

Output:
```
dict_keys(['PostCode', 'LocalAuthorityBusinessID', 'Scores', 'RatingValue', 'BusinessName', 'AddressLine1', 'BusinessType', 'Lng', '_id', 'LocalAuthorityEmailAddress', 'FHRSID', 'LocalAuthorityName', 'BusinessTypeID', 'NewRatingPending', 'RatingKey', 'Geocode', 'SchemeType', 'Region', 'RatingDate', 'LocalAuthorityCode', 'Lat', 'LocalAuthorityWebSite'])
```

<!-- grade_id: extract-data, points: 0, grade: false, solution: true -->
```python
def get_data():
    # YOUR CODE HERE
    cursor = db.uk.find({'Region': 'london', 'RatingValue': {'$ne':None}, 'Geocode':{'$ne':None}},
                    {'BusinessName': 1, 'BusinessType': 1, 'Lat': 1, 'Lng': 1, 
                     'AddressLine1': 1, 'RatingValue':1, '_id': 0}).limit(200)
    
    df = pd.DataFrame(list(cursor))
    return wgs84_to_web_mercator(df, lon="Lng", lat="Lat")

    raise NotImplementedError()

#df = get_data().head()
#type(df)
```

<!-- grade_id: extract-data-test, points: 4.0, grade: true, solution: false -->
```python
# You don't need to write anything here
test_data =  get_data()
assert_equal(type(test_data),pd.DataFrame)

for index, row in test_data.iterrows():
    try:
        i = int(row['RatingValue'])
    except:
        raise AssertionError('There is a row which is not an integer.  '
                             'Make sure you exclude all those without a RatingValue')
        assert False

#print('All tests passed successfully')
assert_equal(len(test_data.index), 200)
print('All tests passed successfully')
```

Output:
```
All tests passed successfully
```

### Question 1(b)  [2 marks]

Create a function **`get_source`** which takes a **`DataFrame`** as a parameter and manipulates it to prepare for addition to the plot.  The function should:

- **Contain a column `Colour`**, which contains a hex string of the colour with which to display the establishment on the map, e.g., #d53e4f.  This should be used to distinguish different RatingValue values of the businesses.
- **RatingValues will be displayed by different colours** using an appropriate palette such as **`Spectral6`** from the standard [Bokeh palettes](http://bokeh.pydata.org/en/latest/docs/reference/palettes.html)
- The function should **accept an integer** as a number to filter the businesses by `RatingValue`.
- If the rating value is equal to **`-1`**, then **all businesses should be included**.  Otherwise, the data should be filtered to only include businesses with a `RatingValue` of the value passed.
- The function should **return** the result as a **`DataFrame`**  

*Hint: Week 5, Guided Exercise 2, Data Sources*

<!-- grade_id: q1b-get-source, points: 0, grade: false, solution: true -->
```python
def get_source(df, data_filter=-1):
    # YOUR CODE HERE
    for var in range(5, -1, -1):
        df.loc[df.RatingValue == var, 'Colour'] = Spectral6[5 - var]
        
    if data_filter > -1:
        df = df.loc[df['RatingValue'] == data_filter]

    return df
    
    raise NotImplementedError()
```

```python
#get_source(get_data(), 3)
```

<!-- grade_id: q1b-get-source-tests, points: 2.0, grade: true, solution: false -->
```python
# You don't need to write anything here
test_source = get_source(get_data())
assert_equal(type(test_source),pd.DataFrame)
# Check the colours are different
rating_value_1 = test_source.loc[test_source['RatingValue'] == 1]
rating_value_2 = test_source.loc[test_source['RatingValue'] == 2]
rating_value_3 = test_source.loc[test_source['RatingValue'] == 3]
rating_value_4 = test_source.loc[test_source['RatingValue'] == 4]
rating_value_5 = test_source.loc[test_source['RatingValue'] == 5]

colour_list = [rating_value_1['Colour'].values[0],
               rating_value_2['Colour'].values[0],
               rating_value_3['Colour'].values[0],
               rating_value_4['Colour'].values[0],
               rating_value_5['Colour'].values[0]]
# If they are all different, then the length of the set should be 5
assert_equal(len(set(colour_list)), 5)

fields = test_source.columns.values
assert 'x' in fields
assert 'y' in fields

# Test that the x and y are the correct web mercator format
test_lng = rating_value_1['Lng'].values[0]
test_lat = rating_value_1['Lat'].values[0]

test_x = rating_value_1['x'].values[0]
test_y = rating_value_1['y'].values[0]

k = 6378137
# print(test_lat * (k * np.pi/180.0))

assert_equal(test_x, test_lng * (k * np.pi/180.0))
assert_equal(test_y, np.log(np.tan((90 + test_lat) * np.pi/360.0)) * k)
print('All tests passed successfully')
```

Output:
```
All tests passed successfully
```

### Question 1(c) [2 marks]

Create a function **`get_map`**, which **returns** a map of London using the [`STAMEN_TERRAIN`](http://bokeh.pydata.org/en/latest/docs/reference/tile_providers.html#bokeh.tile_providers.STAMEN_TERRAIN) tile
 
- The function should **return** a **type** `Figure`
- The figure should include **all** the available **[Pan/Drag tools](http://bokeh.pydata.org/en/latest/docs/user_guide/tools.html#pan-drag-tools)**, the **reset** tool, and the **mouse wheel zoom** tool 
- The map should display London, at an appropriate zoom level for the data
 
*Hint: To get the `x,y` values surrounding London, look at the smallest and largest `x` and `y` values in the data*  
*Hint: Week 4, Guided Exercise 3, Residual Analysis (attributes)*  
*Hint: Week 5, Guided Exercise 2, Data Sources (Bokeh Map Tiling)*

<!-- grade_id: create-plot, points: 2.0, grade: true, solution: true -->
```python
from bokeh.io import output_file, show
from bokeh.layouts import column
from bokeh.plotting import figure


def get_map(data):
    """
    In this function you return a figure with a map background.  The background should be centred
    on London at an appropriate zoom level
    """
    # YOUR CODE HERE
#figure should include all the available Pan/Drag tools, the reset tool, 
#and the mouse wheel zoom
    x_range = min(data['x']), max(data['x'])
    y_range = min(data['y']), max(data['y'])
    tools = 'box_select, box_zoom, lasso_select, pan, resize, wheel_zoom, reset'
    p = figure(tools=tools, x_range=x_range, y_range=y_range)
    #tested: x_range=(min(data['x'], max(data['x'])), y_range=(min(data['y']), max(data['y']))
    p.axis.visible = False
    p.add_tile(STAMEN_TERRAIN)
    return p

    raise NotImplementedError()
show(get_map(get_source(get_data())))
```

### Question 1(d) [3 marks]

Write code which creates and shows a figure `london_map` using the `get_map()` function, obtains a data source for the figure using `get_data()`, then uses the `circle` method to add the data to the map.  

- You should call the output of the `circle` function `data_points`.
- The dots you add to the map should have a size of 10, no border, and a `fill_alpha` of 0.8
- You should call your map `london_map`
- Your code should contain a variable name for the dots added to the map

N.B. You are not being asked to create a function for this question  

*Hint: Week 4, Guided Exercise 3, Fitting a Model - Residual Analysis*  
*Hint: Week 5, Guided Exercise 2, Widgets*

<!-- grade_id: add-data-to-map, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

data = get_data()
source = get_source(data) 
london_map = get_map(source)

#the glyph layer is data_points
data_points = london_map.circle(source= source, x='x', y='y', size=10, fill_color='Colour', line_color = None, 
                                fill_alpha=0.8)
show(london_map, notebook_handle=True)
#raise NotImplementedError()

```

Output:
```
<bokeh.io._CommsHandle at 0x7f89a1bf9400>
```

<!-- grade_id: add-data-to-map-tests, points: 3.0, grade: true, solution: false -->
```python
fig_source = london_map.select(GlyphRenderer)[0].data_source

assert_equal(fig_source.data['x'][0], source['x'][0])
assert_equal(fig_source.data['y'][0], source['y'][0])
assert_equal(fig_source.data['Colour'][0], source['Colour'][0])
#assert_equal(fig_source.data['fill_color'][0], ds.data['Colour'][0])    

glyph = london_map.select(GlyphRenderer)[0].glyph
assert_equal(glyph.line_color, None)
assert_equal(glyph.size, 10)
assert_equal(glyph.fill_alpha, 0.8)
#glyph.fill_color == (ds.data['Colour'][0])
print('All tests passed successfully')
```

Output:
```
All tests passed successfully
```

## Question 2: Make it Interactive

### Question 2(a) [2 marks]

Create a function **`callback`** for later use, which updates the visible businesses on the map according to their **`RatingValue`**.  The function should have parameter **`rating`** which specifies the value to filter by, calling **`get_source`**.  Use the **`source`** variable from Question 1(d) to update the map.  

*Hint: Week 4, Guided Exercise 2, Bokeh Charts*  
*Hint: Week 5, Guided Exercise 2, Data Sources - Widgets (update figure)*

<!-- grade_id: callback-function, points: 2.0, grade: true, solution: true -->
```python
def callback(rating):
    # YOUR CODE HERE
    
    df = get_source(get_data(), rating)
    data_points.data_source.data['x'] = df['x']
    data_points.data_source.data['y'] = df['y']
    data_points.data_source.data['Colour'] = df['Colour']
    
    show(london_map, push_notebook())

    #raise NotImplementedError()

```

```python
#callback(3)
```

### Question 2(b) [2 marks]

Using **`ipywidgets`**, create an [interactive](http://ipywidgets.readthedocs.io/en/latest/examples/Using%20Interact.html) `IntSlider` widget, which calls the **`callback`** function when it updates.

Return the interactive widget in the function **`set_interactive()`**  

*Hint: Week 5, Guided Exercise 2, Widgets*

<!-- grade_id: widget, points: 2.0, grade: false, solution: true -->
```python
def set_interactive():    
    # YOUR CODE HERE
    
    i = interactive(callback, rating=(-1, 5, 1))
    return i

    #raise NotImplementedError()

```

```python
s = set_interactive()
s
```

Output:
```
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
```

<!-- grade_id: widget-tests, points: 2.0, grade: true, solution: false -->
```python
# You don't need to write anything here
old_callback = callback
del callback

try:
    set_interactive()
except NameError as e:
    pass
else:
    raise AssertionError('You have not called the callback function in your code')
finally:
    callback = old_callback
    del old_callback
assert_equal(type(set_interactive()), ipywidgets.widgets.widget_box.Box)
```

## Question 3: Extend the Visualisation

*Applying question 2 solutions*  
  
Now you have created an initial visualisation, you are going to add the following components to it: 
- **Hover** text, so that each dot will give information about the business when you hover
- A **drop down menu** to limit the type of business

You will also be asked to explain a possible use-case for this chart, and offer a suggestion as to how it could be improved.

NOTE: There are discretionary marks available for good visualisation practice

### Question 3(a) [5 marks]

Create a function **`get_hover`**, which returns a [HoverTool](http://bokeh.pydata.org/en/latest/docs/user_guide/tools.html#hover-tool) to be added to the map.  When the cursor hovers over any circle, the following information should be displayed:
- The **name** of the establishment
- The **type** of the establishment
- The **`RatingValue`** of the establishment

Your function should **return** the [HoverTool](http://bokeh.pydata.org/en/latest/docs/user_guide/tools.html#hover-tool).

N.B You may need to read the documentation available to ensure that you match a suitable title for the **field** with the **field value** using **`@`**.  

*Hint: Week 5, Guided Exercise 2, Tools*

```python
data.columns
```

Output:
```
Index(['AddressLine1', 'BusinessName', 'BusinessType', 'Lat', 'Lng',
       'RatingValue', 'x', 'y', 'Colour'],
      dtype='object')
```

<!-- grade_id: hover, points: 5.0, grade: true, solution: true -->
```python
def get_hover():
    # YOUR CODE HERE
    
    hover = HoverTool()
    hover.tooltips = [
    ("Name", "@BusinessName"),
    ("Type", "@BusinessType"),
    ("Rating", "@RatingValue"),
    ]
    return hover
    
    raise NotImplementedError()
```

### Question 3(b) [5 marks]

Using [`interactive`](http://ipywidgets.readthedocs.io/en/latest/examples/Using%20Interact.html), create a **dropdown menu** which allows the user to choose between different **types of business**.  The map should update automatically as you select.

- You should create two functions **`filter_business_types`**, which the dropdown menu will call when it changes.  This function should update the data source
- You will need to obtain a list of the different types of business in the database
- You should **return** an **`interactive`** object, which should call the function **`filter_business_types`**, which updates the **`source`** variable accordingly

N.B. Should you use the original data set?  

*Hint: Week 5, Guided Exercise 2, Data Sources - Tools*

<!-- grade_id: business-types, points: 5.0, grade: true, solution: true -->
```python
def filter_business_types(business_type):
    # YOUR CODE HERE
    source_filtered = data.loc[source['BusinessType'] == business_type]
    dots.data_source.data['x'] = source_filtered['x']
    dots.data_source.data['y'] = source_filtered['y']
    dots.data_source.data['BusinessType'] = source_filtered['BusinessType']
    show(mappy, push_notebook())

def get_dropdown_list():
    # YOUR CODE HERE

    business_list = db.uk.distinct('BusinessType')
    dropdown = interact(filter_business_types, business_type=business_list)
    push_notebook()
    return dropdown
    
    raise NotImplementedError()

# DISPLAY THE MAP
data = get_data()
source = get_source(data)
mappy = get_map(data)
dots = mappy.circle(source=source, x='x', y='y',fill_color='Colour', size=10, fill_alpha=0.8, line_color=None)
mappy.add_tools(get_hover())
dropdown = get_dropdown_list()
slider = set_interactive()
#show(mappy, notebook_handle=True)
```

Output:
```
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
/opt/conda/lib/python3.5/site-packages/bokeh/models/sources.py:81: BokehUserWarning: ColumnDataSource's columns must be of the same length
  lambda: warnings.warn("ColumnDataSource's columns must be of the same length", BokehUserWarning))
```

```python
# You don't have to write anything here
# Display the widgets
HBox([slider, dropdown])
# HBox([slider])
# Why do you think we might get a Bokeh warning diplayed when we use our data?
```

Output:
```
/opt/conda/lib/python3.5/site-packages/ipykernel/jsonutil.py in <listcomp>(.0)
    151 
    152     if isinstance(obj, list):
--> 153         return [json_clean(x) for x in obj]
    154 
    155     if isinstance(obj, dict):

/opt/conda/lib/python3.5/site-packages/ipykernel/jsonutil.py in json_clean(obj)
    171 
    172     # we don't understand it, it's probably an unserializable object
--> 173     raise ValueError("Can't clean for JSON: %r" % obj)

ValueError: Can't clean for JSON: <function filter_business_types at 0x7fe745808598>
```

### Question 3(c) [5 marks]

Describe a use case for which an application like this would be useful, and suggest one way which it could be improved.

<!-- grade_id: visualisation-commentary, points: 5.0, grade: true, solution: true -->
YOUR ANSWER HERE<br>
This application would be very useful on a website similar to Tripadvisor. The aim: to assist the publics food location choices using geospatial intelligence and interactive filters. The main reasons why it would be useful are as follows:
1. The map glyphs for food hygiene ratings are colour coded. This type of retinal variable assists the human brain in dintinguishing groups classified by colour.  This is easier than if the ratings were grouped by shape. The colour groups will also help to identify patterns, such as clustering, in the data.
2. By using Business Type dropdowns and RatingValue sliders, we remove chartjunk from the map. 

One limitation and solution of the application:<br>
Is the web map providing real-time data updates and location intelligence (i.e. where is the customer in relation to the location of a business)? This would be desirable for customers to know that the data visualisation is up-to-date, to aid decision making and navigation. <br>

The solution for real time updates would be to create a tool to extract data via the food hygiene API to the MongoDB as soon as new XML files are added daily. A visualisation reference using bokeh attribution for the most recent date could be added to the map application. And the solution to capture the customers location is to use a mobile geolocation API and create a glyph on the food hygiene map for navigation planning and design. 


