# Data Analysis - Due 05th August (23:59)


In this assessment, we will use the Python libraries `Pandas` and `NumPy` to perform basic operations on the data relating to food hygiene in Wandsworth.  In addition we will also perform linear regression on data about vehicle mileage and price.

Run the cell below each time you load the page to make sure that all the imports are done correctly

```python
import numpy as np
import pandas as pd

from bokeh.io import output_notebook, show
from bokeh.charts import *
output_notebook()

from nose.tools import *

import pymongo
from pymongo import MongoClient
client = MongoClient('mongodb://cpduser:M13pV5woDW@mongodb/health_data', 27017)
db = client.health_data

from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

print('All libraries successfully loaded')

# Create a simple bar chart with made up data
bar_chart = Bar(pd.Series([8,6,2,7,9,3]))
# Display the bar chart
show(bar_chart)
```

Output:
```
All libraries successfully loaded
```

## Question 1: Statistics

### Question 1(a) [4 marks]

Create a **function** `evaluate_random_data` which has the following **parameters**:
- `n` the size of the Pandas `Series` you are to create
- `evaluation` a string should say which task to perform on the `Series` object.
- `seed` an integer which should instantiate a numpy `RandomState` instance, with a default value of 543210

The function should do the following:
- Create an array of size `n` which is full of random numbers integers between 0 and 100 (inclusive)
- The `evaluation` should be able to perform mean, median and standard deviation functions on the created array
- The function should **return** the output of the evaluation on the array.  It should be able to cope with `"mean"`, `"median"` and `"std"` (standard deviation).  

*Hint: Week 4, Guided Exercise 2, Randomness - Importing Data - Bokeh Charts*

<!-- grade_id: random-data-array, points: 0, grade: false, solution: true -->
```python
def evaluate_random_data(n, evaluation, seed=543210):
    # YOUR CODE HERE

    rs = np.random.RandomState(seed)
    listy = []
    for var in range(n):
        listy.append(rs.randint(0, 101))
    
    df = pd.Series(listy)
    
    if evaluation == 'mean':
        return df.mean()
    elif evaluation == 'median':
        return df.median()
    elif evaluation == 'std':
        return df.std()
    else:
        return None
    
    raise NotImplementedError()

```

<!-- grade_id: random-data-array-tests, points: 4, grade: true, solution: false -->
```python
# Check a RandomState object is instantiated
tmp_rs = np.random.RandomState
del np.random.RandomState

try:
    evaluate_random_data(100, 'mean')
except AttributeError:
    print('The function correctly uses RandomState')
else:
    raise AssertionError('The function does not use RandomState')
finally:
    np.random.RandomState = tmp_rs

    
assert_equal(evaluate_random_data(100, 'mean', 511419), 47.39)
assert_equal(evaluate_random_data(100, 'mean', 364235), 48.45)

assert_equal(evaluate_random_data(100, 'median', 511419), 46)
assert_equal(evaluate_random_data(100, 'median', 364235), 44)

output = evaluate_random_data(100, 'std', 511419)
"{0:.4f}".format(round(output,4))
assert_equal("{0:.4f}".format(round(output,4)), '29.2968')

output = evaluate_random_data(100, 'std', 364235)
"{0:.4f}".format(round(output,4))
assert_equal("{0:.4f}".format(round(output,4)), '30.5159')

print('All tests successfully passed')
```

Output:
```
The function correctly uses RandomState
All tests successfully passed
```

### Question 1(b) [3 marks]

The function `get_data` (provided) returns a Pandas `Series` for the given collection.  Using data from this function, and the PyMongo `collection_names` function, write a function `get_means` which takes the mean `RatingValue` from the first `n` collections in the `collection_names()` function.

- You should **sort** `collection_names` in ascending alphabetical order.  Note that this is a Python `list` rather than a `Cursor`, so use **[this guide to sorting lists](https://wiki.python.org/moin/HowTo/Sorting)** to help you
- The function should **return** a **Pandas `Series` object**, with a **name** of **`RatingValueMeans`**.
- You should use the function from Question 1(a) to get a `Series` object of the `RatingValue`, and **obtain the mean** from that object
- N.B. Be very careful you properly exit your loop!

*Hint: Week 3, Guided Exercise 4, Cursors*  
*Hint: Week 4, Guided Exercise 2, Importing Data - Bokeh Charts*

<!-- grade_id: q1a-get-data, points: 0, grade: false, solution: true -->
```python
def get_data(collection):
    cursor = collection.find()
    rating_values = pd.DataFrame(list(cursor))['RatingValue']
    return rating_values

def get_means(n):
    # YOUR CODE HERE
    
    var = 0
    listy = []
    
    for m in sorted(db.collection_names()):
        if var >= n:
            break
        var += 1 
        
        listy.append(get_data(db[m]).mean()) 
        
    return pd.Series(listy, name = 'RatingValueMeans')
    
    raise NotImplementedError()
    
#means = get_means(10)
#means
#type(means)
```

<!-- grade_id: q1-get-data-answer, points: 3, grade: true, solution: false -->
```python
# You don't need to write anything here
means = get_means(10)

assert_equal(type(means), pd.Series)
assert_equal(len(means), 10)
assert_equal(round(means.sum(), 4), 27.2436)
means = get_means(12)
assert_equal(len(means), 12)
assert_equal(round(means.sum(), 4), 36.2604)
print('All tests passed successfully')
```

Output:
```
All tests passed successfully
```

### Question 1(c) [4 marks]

Create a function `get_sample_mean_distribution` which **returns** a series of the distribution of sample means of the `RatingValue` of the data from a given series.  The function should be defined as follows:
- `data`: A `Series` object of the mean ratings of establishments 
- `n`: integer, the size of the sample
- `m=1000`: integer, the amount of times to repeat the procedure
- `seed=543210`: integer, default value 543210.  This should set a [`RandomState`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.random.RandomState.html) instance to be initialised with this value as an argument.
- You should use the Pandas [`sample`](http://pandas.pydata.org/pandas-docs/version/0.19.2/generated/pandas.DataFrame.sample.html) function to generate your samples.

The function should **return** the output as a **Pandas `Series` object**.  

*Hint: Week 4, Guided Exercise 2, Randomness*

```python
# code test cell
listy = []
m = 1000
n = 20
seed = 543210
var = np.random.RandomState(seed)
data = pd.Series([5, 5, 5, 5, 5, 4, 5, 4, 5, 4, 5, 5, 3, 4, 4, 5, 5, 5, 4, 5, 5, 5, 5, 4, 2, 4, 5, 5, 4, 5])
count = 0
#30 ints in the series

#calculate the mean of the sample m times return a list

while count < m:
    samp = data.sample(n, random_state=var).mean()
    listy.append(samp)
    count += 1

#print(pd.Series(listy))

```

<!-- grade_id: sample-mean-distribution, points: 0, grade: false, solution: true -->
```python
def get_sample_mean_distribution(data, n, m=1000, seed=543210):
    # YOUR CODE HERE
    
    var = np.random.RandomState(seed)
    listy = []
    count = 0
    #calculate the mean of the sample m times return a list
    while count < m:
        samp = data.sample(n, random_state=var).mean()
        listy.append(samp)
        count += 1
    
    return pd.Series(listy)
        
    raise NotImplementedError()

#data = pd.Series([5, 5, 5, 5, 5, 4, 5, 4, 5, 4, 5, 5, 3, 4, 4, 5, 5, 5, 4, 5, 5, 5, 5, 4, 2, 4, 5, 5, 4, 5])    
#get_sample_mean_distribution(data, 20, 1000, 123456)
```

<!-- grade_id: sample-mean-distribution-tests, points: 4, grade: true, solution: false -->
```python
# You don't need to write anything here
m = 1000
n = 20
data = pd.Series([5, 5, 5, 5, 5, 4, 5, 4, 5, 4, 5, 5, 3, 4, 4, 5, 5, 5, 4, 5, 5, 5, 5, 4, 2, 4, 5, 5, 4, 5])
seed = 123456
actual_output = get_sample_mean_distribution(data, n, m, seed)

assert_equal(len(actual_output), 1000)
assert_equal(round(actual_output.sum(),4), 4528.6)

# Check randomness is working:
# The same seed should lead to the same result
seed_test_equal = get_sample_mean_distribution(data, n, m, seed)
assert_equal(actual_output.mean(), seed_test_equal.mean())

# A different seed should not be equal
seed_test_not_equal = get_sample_mean_distribution(data, n, m, 54321)
assert_not_equal(actual_output.mean(), seed_test_not_equal.mean())

print('All tests successfully passed')
```

Output:
```
All tests successfully passed
```

### Question 1(d) [2 marks]

Using the output from Question 1(c), **plot a histogram** illustrating the distribution, and **return** the histogram as the output of the function `plot_sample_mean_distribution`.  

This function should take the same parameters as the `get_sample_mean_distribution`, so that it can call that function to obtain the data.  The histogram should have **10 bins**.  

*Hint: Week 1, Guided Exercise 1, Python Primer*  
*Hint: Week 2, Guided Exercise 2, Python for Data Science*  
*Hint: Week 4, Guided Exercise 2, IQR & Outliers - Bokeh Charts*  
*Hint: Week 4, Guided Exercise 3, Residual Analysis*

<!-- grade_id: sample-mean-distribution-plot, points: 0, grade: false, solution: true -->
```python
def plot_sample_mean_distribution(data, n, m=1000, seed=543210):
    # YOUR CODE HERE
  
    var = get_sample_mean_distribution(data, n)
    hist = Histogram(var , bins=10)
    show(hist)
    return hist
    
    raise NotImplementedError()
    
#data = pd.Series([5, 5, 5, 5, 5, 4, 5, 4, 5, 4, 5, 5, 3, 4, 4, 5, 5, 5, 4, 5, 5, 5, 5, 4, 2, 4, 5, 5, 4, 5])
#plot_sample_mean_distribution(data, 20)
```

<!-- grade_id: sample-mean-distribution-plot-tests, points: 2, grade: true, solution: false -->
```python
# You don't need to write anything here
data = pd.Series([5, 5, 5, 5, 5, 4, 5, 4, 5, 4, 5, 5, 3, 4, 4, 5, 5, 5, 4, 5, 5, 5, 5, 4, 2, 4, 5, 5, 4, 5])

# assert_equal())
assert_equal(type(plot_sample_mean_distribution(data, 20)), Chart)
# Check Histogram
old_hist = Histogram
del Histogram

try:
    plot_sample_mean_distribution(data, 20)
except NameError:
    pass
else:
    raise AssertionError("The chart does not appear to be a Histogram")
finally:
    Histogram = old_hist
    del old_hist

print('All tests successfully passed')
```

Output:
```
All tests successfully passed
```

## Question 2: Linear Regression

This question uses a different dataset in order to illustrate linear regression.  We'll use a dataset similar to the cars dataset given as an example in the videos.

### Question 2(a) [1 mark]

Create a function `get_car_data` which uses the Pandas `read_csv` to import the data in the `cars.csv` file and **return** a `DataFrame` of the data (The `cars.csv` file is in the same directory in the notebook)  

*Hint: Week 4, Guided Exercise 3, Classification - Importing Data*

<!-- grade_id: cars-loading, points: 0, grade: false, solution: true -->
```python
def get_car_data():
    # YOUR CODE HERE
    
    data = pd.read_csv('cars.csv')
    return data

    raise NotImplementedError()

```

<!-- grade_id: cars-loading-test, points: 1, grade: true, solution: false -->
```python
# You don't need to write anything here
assert_equal(len(get_car_data()), 1155)
assert_equal(type(get_car_data()), pd.DataFrame)
print('All tests successfully passed')
```

Output:
```
All tests successfully passed
```

### Question 2(b) [3 marks]

Create a **function** `remove_outliers` which removes the outliers from the dataset according to the $lower - (1.5 * IQR)$ and $upper + (1.5 * IQR)$ rule.  The function should have the following parameters:

- `data` a Pandas `DataFrame`
- `field` a string, which will say for which of the fields in the `DataFrame` to locate and remove outliers  

The function should **return** the `DataFrame` without the outliers.

*Hint: Week 4, Guided Exercise 2, IQR*  
*Hint: Week 4, Guided Exercise 3, Residual Analysis*

```python
# code test cell
data = get_car_data()
data = pd.DataFrame(data)
isinstance(data, pd.DataFrame)
#data.head()

#first determine the upper and lower value to subset out
upper_var = data['Price'].quantile(0.75)
lower_var = data['Price'].quantile(0.25)
iqr = upper_var - lower_var
min_val = lower_var - (1.5 * iqr)
max_val = upper_var + (1.5 * iqr)
#now create a new dataset removing the min and max vals
new_data = data.loc[data['Price'] <= max_val]
new_data = new_data.loc[data['Price'] >= min_val]
#min_val = -2255
#max_val = 12545
new_data.shape

```

Output:
```
(1101, 5)
```

<!-- grade_id: remove-outliers, points: 0, grade: false, solution: true -->
```python
def remove_outliers(data, field):
    # YOUR CODE HERE
    
    data = pd.DataFrame(data)

    #first determine the upper and lower quantile values to subset out
    upper_var = data[field].quantile(0.75)
    lower_var = data[field].quantile(0.25)
    iqr = upper_var - lower_var
    min_val = lower_var - (1.5 * iqr)
    max_val = upper_var + (1.5 * iqr)
    
    #now create a new dataset removing the min and max vals
    new_data = data.loc[data[field] <= max_val]
    new_data = new_data.loc[data[field] >= min_val]
    
    return new_data

    raise NotImplementedError()
```

<!-- grade_id: remove-outliers-tests, points: 3, grade: true, solution: false -->
```python
# You don't have to write anything here
data = get_car_data()

data = remove_outliers(data, 'Price')
assert_equal(data.count()['Price'], 1101)

data = remove_outliers(data, 'Mileage')
assert_equal(data['Mileage'].count(), 1100)

data = remove_outliers(data, 'Year')
assert_equal(data['Year'].count(), 1092)
print('All tests successfully passed')
```

Output:
```
All tests successfully passed
```

### Question 2(c) [2 marks]
Create a function `scatter_feature`, which generates a **scatter** plot of data.  The function should take two parameters: `data`, a `DataFrame` of the regression data, and `feature_name`, the column name of the feature to be used as a predictor variable.  The following should also be applied:
- The **response** variable is the `Price` columnn
- The function should return a **`Chart`** instance of the scatter plot.  

*Hint: Week 4, Guided Exercise 2, Bokeh Charts*

<!-- grade_id: scatter-features, points: 1, grade: true, solution: true -->
```python
def scatter_feature(data, feature_name):
    """
    This function returns a scatter plot for the particular `feature_name` v Price
    `data` should be a `DataFrame`
    `feature_name` should be a string
    """
    # YOUR CODE HERE
    
    scatter = Scatter(data, x=feature_name, y='Price')
    #show(scatter) #show is used when calling on the scatter_feature function
    return scatter
    
    #raise NotImplementedError()
show(scatter_feature(data, 'Mileage'))
show(scatter_feature(data, 'Year'))
```



### Question 2(d) [3 marks]

Using the `scikit-learn` libarary, write a function `split_data` which generates a **training** and a **test** dataset, based on features passed in as arguments.  The function should have three parameters:

- **`data`** a `DataFrame` containing the data to be evaluated
- **`features`** a list of strings which will be used to select the features from the dataset
- **`seed`** an integer with default value `543210`, which  should be used to initialise a `RandomState` (as with Question 1(a), 
- The function should return the output of **[`train_test_split`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html)**, using that `RandomState`, and data with the selected features from the `features` parameter. 
- The **test size** of the returned data should be **20%**  
- As with question 2c the **response** variable shuld be the **'Price'** columnn  

*Hint: Week 4, Guided Exercise 3, Training and Testing Data*

<!-- grade_id: scatter-data, points: 0, grade: false, solution: true -->
```python
def split_data(data, features, seed=543210):
    # YOUR CODE HERE
    #used the module sklearn and function train_test_split
    
    X = data[features]
    y = data['Price'] 
    return train_test_split(X, y, test_size=0.2, random_state=np.random.RandomState(seed=seed))
    raise NotImplementedError()

#x1, x2, y1, y2 = split_data(data, ['Mileage'], 123456)
#x1.index[0]
```

<!-- grade_id: scatter-data-tests, points: 3, grade: true, solution: false -->
```python
# You don't need to write anything here
data = get_car_data()

x1, x2, y1, y2 = split_data(data, ['Mileage'])
assert_equal(len(x1), 924)
assert_equal(len(x2), 231)

assert_equal(x1.index[0], 607)
x1, x2, y1, y2 = split_data(data, ['Mileage'], 123456)
assert_equal(x1.index[0], 642)
print("All tests successfully passed")
```

Output:
```
All tests successfully passed
```

### Question 2(e) [4 marks]

Create a function `evaluate_model` to **fit** and **evaluate** the linear regression model using the test the accuracy of the model with different features using $R^2$.  The function should contain the same parameters as `split_data`, so it can call that function.

The model should be evaluated based on the test data.  

*Hint: Week 2, Guided Exercise 2, Exercise 3a*  
*Hint: Week 4, Guided Exercise 3, Lineaar Regression - Fitting The Model - Residual Analysis*

```python
lm = LinearRegression()
X_train, X_test, y_train, y_test = split_data(data, ['Year'], 511419)
model = lm.fit(X_train, y_train)
score = lm.score(X_test, y_test)
print(model, score)
print(round(score, 4))
```

Output:
```
LinearRegression(copy_X=True, fit_intercept=True, n_jobs=1, normalize=False) 0.252556080999
0.2526
```

<!-- grade_id: predict-model, points: 0, grade: false, solution: true -->
```python
def evaluate_model(data, features, seed=543210):
    # YOUR CODE HERE
    
    lm = LinearRegression()
    X_train, X_test, y_train, y_test = split_data(data, features, seed)
    model = lm.fit(X_train, y_train)
    score = lm.score(X_test, y_test)
    return score
    
    raise NotImplementedError()
```

<!-- grade_id: predict-model-tests, points: 4, grade: true, solution: false -->
```python
## You don't need to write anything here

data = get_car_data()

assert_equal(round(evaluate_model(data, ['Mileage'], 511419), 4),0.4742)
assert_equal(round(evaluate_model(data, ['Year'], 511419), 4), 0.2526)
assert_equal(round(evaluate_model(data, ['Year', 'Mileage'], 511419), 4), 0.4754)
print("All tests successfully passed")
```

Output:
```
All tests successfully passed
```